# Homework 1

The maximum score of this homework is 100+10 points. Grading is listed in this table:

| Grade | Score range |
| --- | --- |
| 5 | 85+ |
| 4 | 70-84 |
| 3 | 55-69 |
| 2 | 40-54 |
| 1 | 0-39 |

Most exercises include tests which should pass if your solution is correct.
However successful test do not guarantee that your solution is correct.
The homework is partially autograded using many hidden tests.
Test cells cannot be modified and empty cells cannot be deleted.

Your solution should replace placeholder lines such as:

    ### YOUR CODE HERE
    raise NotImplementedError()

Please do not add new cells, they will be ignored by the autograder.

**VERY IMPORTANT** Before submitting your solution (pushing to the git repo),
run your notebook with `Kernel -> Restart & Run All` and make sure it
runs without exceptions.

## Submission

GitHub Classroom will accept your last pushed version before the deadline.
You do not need to send the homework to the instructor.

## Plagiarism

When preparing their homework, students are reminded to pay special attention to Title 32, Sections 92-93 of Code of Studies (quoted below). Any content from external sources must be stated in the students own words AND accompanied by citations. Copying and pasting from an external source should be avoided and any text copied must be placed between quotation marks. Reports that violate these rules cannot receive a passing grade.

"**Section 92**

(1) The works of another person will be used as follows: a) if a work of another person is used in whole or in part (e.g. by copying, citation, translation from another language or presentation), the source and the name of the author will be indicated if this name is included in the source or – in case of orally presented works – may be clearly identified; b) the work of another person or any part of that will be used – up to a quantity reasonably corresponding to the nature and purpose of the student work – identified as quotations.

(2) Instructors are entitled to review compliance with requirements in this article with computer programmes and databases.

(3) The use of works of another person and the acknowledgement of use will be governed by applicable laws and the relevant rules of the specific discipline.

**Section 93**

(1) If a student fails to meet rules regarding use of works of another person in whole or in part, the student work will be considered as not assessable and the student will not be allowed to obtain the credit of the concerned subject in the specific term.

(2) It will be deemed a disciplinary offence if a student – in breach of the rules regarding use of works of another person – submits or presents a work of another person fully or in a significant part verbatim (word for word) or in terms of its basic concepts or the combined version of several works of another person(s) as their own work.

(3) Based on subsection (1) of Section 52/A. of the Higher Education Act, compliance with the rules regarding the use of works of another person in a master thesis may be reviewed up to five years following the issue of the degree certificate. In case of violation of the above rules, section 52/A of the Higher Education Act will apply."

(BME Code of Studies, p.50)

## 1. Modified Levenshtein distance (20 points)

Standard Levenshtein distance assigns an integer edit distance to any two strings measuring how difficult it would be to turn one string into the other. See [Wikipedia](https://en.wikipedia.org/wiki/Levenshtein_distance).

Create a modify version of Levenshtein distance which discounts letters
that are close to each other on the English keyboard. The keyboard variable below
contains the English keyboard organized into a table.
Two letters are considered close to each other if both their row
and column distance is at most 1. Close keys are at distance 0.5,
others are at distance 1.
This table lists a few examples:

| | | distance |
| ---- | ---- |
| q | w | 0.5 |
| q | e | 1 |
| s | w | 0.5 |
| f | t | 0.5 |
| f | y | 1 |
| f | f | 0 |

Any letter outside the lowercase English alphabet (see the `keyboard` variable below)
is not considered close and you do not need to discount them.

```python
keyboard = [
    ['q', 'w', 'e', 'r', 't', 'y', 'u', 'i', 'o', 'p'],
    ['a', 's', 'd', 'f', 'g', 'h', 'j', 'k', 'l'],
    ['z', 'x', 'c', 'v', 'b', 'n', 'm'],
]
keyboard_mapping = {}
for i, row in enumerate(keyboard):
    for j, c in enumerate(row):
        keyboard_mapping[c] = (i, j)
```

<!-- grade_id: cell-4aabc448be0f8eb8, points: 0, grade: false, solution: true -->
```python
def keyboard_edit_distance(str1, str2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-92748a4d4ec4b9a3, points: 10, grade: true, solution: false -->
```python
assert keyboard_edit_distance("abc", "abc") == 0
assert keyboard_edit_distance("abc", "ab") == 1
```

<!-- grade_id: cell-d6d5e8aeb4f137cb, points: 10, grade: true, solution: false -->
```python
assert keyboard_edit_distance("a", "s") == 0.5
```

## 2. Replace rare words (10 points)

Write a function that takes a text and a number $N$ as parameters and replaces every word other than the most common $N$ in the text with a common symbol. The symbol by default is `__RARE__` but it can be redefined.

Your code should split on spaces only.

You can derive the function definition from the tests.

<!-- grade_id: cell-3f03df0e30541457, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ad7519e408fad885, points: 3, grade: true, solution: false -->
```python
assert replace_rare_words("a b a a a b b b c d d d", 2) == "a b a a a b b b __RARE__ __RARE__ __RARE__ __RARE__"
assert replace_rare_words("a b a b b c", 2, rare_symbol="rare") == "a b a b b rare"
```

<!-- grade_id: cell-03ebc98fa36b7e93, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0454e3a47008a258, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-07ce23e88f73124a, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b83b7f41c5ed954d, points: 2, grade: true, solution: false -->
```python

```

## 3. MutableString (30 points)

Python strings are immutable. Create a mutable string class. The internal representation should be mutable too.

Implement the following features (see the tests below).

- initialization from `str`,
- initialization from text file (loads the file's content into the string),
- assignment (i.e. modifying a character),
  - if the index is out of range, it should fill the blanks with spaces (see the tests below)
- conversion to built-in `str` and `list`. The latter is a list of the characters.
- addition with other `MutableString` instances and built-in strings,
- multiplication with integers. Multiplying a string with 3 means repeating the string 3 times.
- built-in `len` function,
- comparision with strings,
- substring containment with both built-in strings and other MutableString objects,
- in-place upper and lowercasing,
- shallow copying,
- iteration.

Please read all of the tests before writing your solution.

<!-- grade_id: cell-97c9229b87c830f9, points: 0, grade: false, solution: true -->
```python
class MutableString(object):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Briefly describe what internal representation you chose and why you chose it.

What other possibilities did you think about?

(double-click on the text to edit it)

<!-- grade_id: cell-f17f91b5fa516d35, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

### MutableString tests

<!-- grade_id: cell-02165ab6b6c46b4f, points: 2, grade: true, solution: false -->
```python
# initialization
m = MutableString()
assert m == "" and len(m) == 0
```

<!-- grade_id: cell-8cc3f5587430f951, points: 1, grade: true, solution: false -->
```python
# initialization from file
import tempfile
import os

with tempfile.NamedTemporaryFile(mode="w+b", delete=False) as tmpfile:
    tmpfile.write("abc".encode("utf8"))
m = MutableString.from_file(tmpfile.name)
assert m == "abc"
```

<!-- grade_id: cell-8cc3f5587430f953, points: 1, grade: true, solution: false -->
```python
os.remove(tmpfile.name)
```

<!-- grade_id: cell-8cc3f5587430f952, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-9a96e6dcdc4b2fcc, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-81efd81e7a00554e, points: 1, grade: true, solution: false -->
```python
# iteration
m = MutableString("abc")
for c in m:
    print(c)
```

<!-- grade_id: cell-dbb2ff81492d3f72, points: 2, grade: true, solution: false -->
```python
# comparison
m1 = MutableString("abc")
assert m1 == "abc"
m2 = MutableString("abc")
assert id(m1) != id(m2)
assert m1 == m2
```

<!-- grade_id: cell-07a85b37bbb4564e, points: 2, grade: true, solution: false -->
```python
# item setting tests
m1 = MutableString("abc")
assert m1[0] == "a" and m1[2] == "c"
m1[1] = "d"
assert(m1[1] == "d")
```

<!-- grade_id: cell-11a3b1af61d04515, points: 2, grade: true, solution: false -->
```python
# slicing
m1 = MutableString("abc")
assert m1[1:] == "bc"
m1[2:4] = "ad"
assert m1 == "abad"
```

<!-- grade_id: cell-8f9260e50cc39598, points: 2, grade: true, solution: false -->
```python
# conversions
m1 = MutableString("abc d")
assert(list(m1) == list("abc d"))
assert(str(m1) == "abc d")
```

<!-- grade_id: cell-dd0bda8e1c531bba, points: 2, grade: true, solution: false -->
```python
# concatenation
m1 = MutableString("abc")
m2 = MutableString("def")
assert m1 + m2 == "abcdef"
assert m2 + m1 == "defabc"
```

<!-- grade_id: cell-d4e5abf63477846d, points: 1, grade: true, solution: false -->
```python
# multiplication
m1 = MutableString("abc")
m3 = m1 * 3
assert m3 == "abcabcabc"
```

<!-- grade_id: cell-7ae249bf738b8d6a, points: 2, grade: true, solution: false -->
```python
# operator=
m1 = MutableString("abc")
m2 = m1
m2[0] = "A"
assert m1[1:] == "bc"
```

<!-- grade_id: cell-c7cae6bb8bb14f11, points: 2, grade: true, solution: false -->
```python
# copy
from copy import copy

m1 = MutableString("abc")
m2 = copy(m1)
m2[0] = "A"
assert m2 == "Abc"
```

<!-- grade_id: cell-0e1676338f39c777, points: 2, grade: true, solution: false -->
```python
# concatenation with strings
m1 = MutableString("abc")
m2 = m1 + "def"
assert m2 == "abcdef"
m3 = "def" + m1
assert m3 == "defabc"
```

<!-- grade_id: cell-2a01c20eb0bd11fc, points: 2, grade: true, solution: false -->
```python
# in place lowercasing and uppercasing
m1 = MutableString("aBc")
m1.to_upper()
assert m1 == "ABC"
m1 = MutableString("aBc")
m1.to_lower()
```

<!-- grade_id: cell-d8f0ef82c2479b20, points: 1, grade: true, solution: false -->
```python
# containment test
m1 = MutableString("abcdef")
assert "bcd" in m1
m2 = MutableString("bcd")
assert m2 in m1
```

# Text generation

## 3.1 (Same as a laboratory exercise) Write a function that computes N-gram frequencies in a string. (0 point)

<!-- grade_id: cell-e28c0d3638d63105, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c1ff36135c46cfed, points: 0, grade: true, solution: false -->
```python
assert(count_ngram_freqs("abcc", 1) == {"a": 1, "b": 1, "c": 2})
assert(count_ngram_freqs("abccab", 2) == {"ab": 2, "bc": 1, "cc": 1, "ca": 1})
```

## 3.2 Define a random text generator function (20 points).

The function takes 4 arguments:

1. starting text (at least $N-1$ long),
2. target length: length of the output string,
3. n-gram frequency dictionary,
4. N, length of the n-grams.

The function generates one character at a time given the last $N-1$ characters.
The probability of `c` being generated after `ab` is defined as:

$$
P(c | a b ) = \frac{\text{freq}(a b c)}{\text{freq}(a b)},
$$

where  $\text{freq}(a b c)$ is obtained by counting how many times `abc` occurs in the training corpus (`count_ngram_freqs` function).

If the generated text ends with a $N-1$-gram that does not occur in the training data, generate the next character from the full character or ngram distribution.

<!-- grade_id: cell-ff2fea718aeafac3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8ce5a7677377c6fb, points: 20, grade: true, solution: false -->
```python
toy_freqs = count_ngram_freqs("abcabcda", 3)
gen = generate_text("abc", 5, toy_freqs, 3)

assert(len(gen) == 5)
assert(set(gen) <= set("abcd"))
```

## 3.3 Test your solution on a small Wikipedia corpus (10 points).

Collect a sample of at least 1 million characters from Wikipedia using the wikipedia module.

<!-- grade_id: cell-1754f5d1a69a2abf, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## \*3.4 Smoothing (extra exercise, 10 points)

Implement one or more smoothing methods such as Jelinek-Mercer smoothing or Katz's backoff. Train it on your Wikipedia corpus and generate an example.

https://nlp.stanford.edu/~wcmac/papers/20050421-smoothing-tutorial.pdf

<!-- grade_id: cell-163b94be41cd70f2, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Code cleanness and PEP8 (10 points)

This cell is here for technical reasons, you will receive feedback on your code quality here. You do not need to write anything here.

<!-- grade_id: cell-937470ef5cbde557, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE